# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
from pathlib import Path
HOME = str(Path.home())

ALLERGIES_FILE="/data/allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def allergy_count(json_file):
    """(JSON file as str) -> int
    This function finds the number of allergies listed within the file allergies.json. It returns the numbers of allergies as an
    integer value.
    
    >>> allergy_count(ALLERGIES_FILE)
    4
    """
    
    with open(json_file) as f: # opens the json file passed through the parameter # set up to close file automatically
        bundle = json.load(f) # reads json file into a dictionary called bundle
        
    allergy_counter = 0 # set allergy counter variable to zero

    entries = bundle.get('entry') # assigns values in entry key to a list called entries
    
    for entry in entries: # iterates through each entry 
        resource = entry.get('resource') # assigns key-value pairs in resource key to a dict called resource
        
        if resource['resourceType'] == 'AllergyIntolerance' and resource['status'] == 'confirmed': # confirms if the resource is a confirmed allergy
            allergy_counter += 1 # increases allergy counter by 1
            
    return int(allergy_counter) # return allergy counter variable set as an integer

### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(allergy_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_count(ALLERGIES_FILE)
Expecting:
    4
ok
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
### BEGIN SOLUTION

def patient_count(json_file):
    """(JSON file as str) -> int
    This function finds the number of unique patients listed within the file allergies.json. It returns the number of unique patients as an
    integer value.
    
    >>> patient_count(ALLERGIES_FILE)
    2
    """
    
    with open(json_file) as f: # opens the json file passed through the parameter # set up to close file automatically
        bundle = json.load(f) # reads json file into a dictionary called bundle
        
    unq_pt_counter = 0 # set unique patient counter variable to zero
    unq_pt_list = [] # set list of unique patients as empty

    entries = bundle.get('entry') # assigns values in entry key to a list called entries
    
    for entry in entries: # iterates through each entry 
        resource = entry.get('resource') # assigns key-value pairs in resource key to a dict called resource
        patient = resource['patient'] # assigns value in resource key 'patient' to patient
        display = patient['display'] # assigns value in patient key 'display' to display
        
        if display not in unq_pt_list: # if the patient's name is not in the list...
            unq_pt_list.append(display) # ...then add the patient's name to the list
            unq_pt_counter += 1 # increases unique patient counter by 1 if a new patient is added to the unique patient list 
            
        else: # if the patient's name is in the list...
            pass # ...then skip adding to the list and increasing the counter
            
    return int(unq_pt_counter) # return unique patient counter variable set as an integer

### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(patient_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count(ALLERGIES_FILE)
Expecting:
    2
ok
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
### BEGIN SOLUTION

def allergy_per_patient(json_file):
    """(JSON file as str) -> printed str
    This function finds the number of allergies per patient listed within the file allergies.json.
    It prints the patient's name and the number of allergies they have.
    
    >>> allergy_per_patient(ALLERGIES_FILE)
    Jason Argonaut's allergy count: 3.
    Paul Boal's allergy count: 1.
    """
    
    with open(json_file) as f: # opens the json file passed through the parameter # set up to close 
        bundle = json.load(f) # reads json file into a dictionary called bundle
    
    unq_pt_dict = {} # sets dict of unique patients as empty

    entries = bundle.get('entry') #  assigns values in entry key to a list called entries
    
    for entry in entries: # iterates through each entry 
        resource = entry.get('resource') # assigns key-value pairs in resource key to a dict called resource
        
        patient = resource['patient'] # assigns value in resource key 'patient' to patient
        display = patient['display'] # assigns value in patient key 'display' to display
        
        substance = resource['substance'] # assigns value in resource key 'substance' to subtance
        text = substance['text'] # assigns value in substance key 'text' to text
        
        if display not in unq_pt_dict: # if the patient's name is not a key in unqiue patient dict...
            unq_pt_dict[display] = [text] # ...then set the current substance value as a list to the patient's name in the dict
            
        else: # if the patient's name is already a key in unqiue patient dict...
            unq_pt_dict[display].append(text) # ...then add the current substance to the list
    
    for unq_pt in unq_pt_dict.items(): # iterates through each unique patient
        print(f"{unq_pt[0]}'s allergy count: {len(unq_pt[1])}.") # prints the current patient's name and their allergy count
    
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(allergy_per_patient, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_per_patient(ALLERGIES_FILE)
Expecting:
    Jason Argonaut's allergy count: 3.
    Paul Boal's allergy count: 1.
ok
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
Jason Argonaut's allergy count: 3.
Paul Boal's allergy count: 1.
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION

def allergy_list (json_file):
    """(JSON file as str) -> list
    This function finds a patient's name, allergy, and reaction from the file allergies.json. 
    It returns a list of these values which does not include any duplicate name/allergy pairs (meaning only the first reaction found will be provided).
    
    >>> allergy_list(ALLERGIES_FILE)
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Paul Boal', 'PENICILLIN G', 'Bruising'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
    """
    
    with open(json_file) as f: # opens the json file passed through the parameter # set up to close 
        bundle = json.load(f) # reads json file into a dictionary called bundle
    
    allergy_list = [] # sets list of allergies as empty

    entries = bundle.get('entry') # assigns values in entry key to a list called entries

    
    for entry in entries: # iterates through each entry
        resource = entry.get('resource') # assigns key-value pairs in resource key to a dict called resource
        
        patient = resource['patient'] # assigns value in resource key 'patient' to patient
        display = patient['display'] # assigns value in patient key 'display' to display
        
        substance = resource['substance'] # assigns value in resource key 'substance' to subtance
        sub_text = substance['text'] # assigns value in substance key 'text' to sub_text

        reactions = resource['reaction'] # assigns value in resource key 'reaction' to reactions
        
        for reaction in reactions: # iterates through each reaction
            manifestations = reaction.get('manifestation') # assigns key-value pairs in manifestation key to manifestations
            
            for manifestation in manifestations: # iterates through each manifestation
                allergy_text = manifestation.get('text') # assigns value in manifestation key 'text' to allergy_text
                allergy_list.append([display, sub_text, allergy_text]) # adds list consisting of patient's name, allergy, and reaction to allergy list
                break  # skips next manifestation(s) of the current patient and allergy
                  
    return allergy_list # returns the allergy list 
            
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(allergy_list, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_list(ALLERGIES_FILE)
Expecting:
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Paul Boal', 'PENICILLIN G', 'Bruising'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
ok
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list(ALLERGIES_FILE) == output

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION

def allergy_reaction(json_file, patient, substance):
    """(JSON file as str, str, str) -> str
    This function provides the type of reaction of an allergy for a patient. 
    It calls upon the allergy_list function to find the list of patients and their allergies/reactions from the file allergies.json.
    It returns the reaction as a string if the patient and substance parameters are found in the list mentioned above.
    
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') 
    'Itching'
    
    >>> allergy_reaction(ALLERGIES_FILE, 'Helen Troy', 'SHELLFISH-DERIVED PRODUCTS')
    
    >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'BLUEBERRY')
    """
    
    allergies = allergy_list(json_file) # calls on allergy_list function to assign list of patient's name/allergy/reaction to allergies
    
    allergy_reaction = None # sets allergy_reaction to None (nothing)
    
    for entry in allergies: # iterates through each entry
        if entry[0] == patient and entry[1] == substance: # if the first value (name) in the current entry list matches the patient parameter and the second value (allergy) matches the substance parameter...
            allergy_reaction = entry[2] # ...then the third value (reaction) is assigned to allergy_reaction
            
        else: # if either the first or second values do not match...
            pass # ...keep allergy_reaction set to None and go on to the next entry list
       
    return allergy_reaction # returns the reaction to the patient's allergy

### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(allergy_reaction, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') 
Expecting:
    'Itching'
ok
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Helen Troy', 'SHELLFISH-DERIVED PRODUCTS')
Expecting nothing
ok
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'BLUEBERRY')
Expecting nothing
ok
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---
---

# Stretch (Extra) Problem

Work on either of the stretch problems below can earn you up to 25 free points toward the midterm assignment.  That is, if you complete one of these extra problems successfully, you can skip 1 of the problems that will appear on the midterm exam coming up next week.

The midterm will be distributed 10/14 and due 10/24.



---
---

### STRETCH for October 2022 - For those looking for an additional challenge

As I've mentioned in class, CMS is now enforcing a rule around price transparency.  Every facility that take Medicare payments is required to publish a "machine readable" file with it's pricing infomration for a number of common procedures across all of the payers they work with.  There are two examples of such files in the `/data/` directory: `whiteriver.json` and `saline.xml`.

If you want to compare contracted prices across these two hospitals, you'll need to read in the information from both of those files into some kind of data structure, then merge the data together from those two files.  See what you can do.

See if you can create an output file that has the following fields:
* HOSPITAL
* PROCEDURE_CODE
* PAYER
* AMOUNT

If you choose to work on this, you may get stuck at some point and you won't know if you're _doing it right_. Make some assumptions. Document your questions in this notebook.



```
Procedure Code |  Description  |  Gross Charges  |  Aetna  |  QualChoice
```

```python
# WHITERIVER JSON FILE (TEST CODE - DO NOT RUN)
import json

with open('/data/whiteriver.json') as j:
    root = json.load(j)

s = open('wr_test.txt', 'w')

s.write(root['root']['HospitalorFacilityName'] + '\n\n')

wr_charges = root['root']['StandardCharges']

for wr_charge in wr_charges:
    s.write('\tProcedureCode: ' + wr_charge['ProcedureCode'] + '\n')
    s.write('\t\tDescription: ' + wr_charge['Description'] + '\n')
    s.write('\t\tGross Charges: ' + wr_charge.get('OutpatientGrossCharge') + '\n')
    s.write('\t\tAetna Pricing: ' + wr_charge.get('AETNA_Outpatient') + '\n')
    s.write('\t\tQualChoice Pricing: ' + wr_charge.get('QUALCHOICE_Outpatient') + '\n\n')
    
s.close()
```

```python
# SALINE XML FILE (TEST CODE - DO NOT RUN)
from lxml import objectify

with open('/data/saline.xml') as x:
    sa = objectify.parse(x)

s = open('sa_test.txt', 'a')
test = sa.getroot()

for facility in test.findall('Facility'):
    name = facility.get('Name')
    s.write(name.title() + '\n\n')
    
patients = facility.getchildren()

for patient in patients:
    if patient.get('Type') == 'Outpatient':
        #outpatient = patient.get('Type')
        sa_charges = patient.getchildren()

for sa_charge in sa_charges:
    codes = sa_charge.getchildren()

for code in codes:
    if code.get('Code') < '19082':
        s.write('\tProcedure Code: ' + code.get('Code') + '\n')   
        items = code.getchildren()

        s.write('\t\tDescription: ' + items[0] + '\n')
        s.write('\t\tGross Charges: ' + str(items[2]) + '\n')

        for item in code.iter('Contract'):
            if item.get('Payer') == 'AETNA - (POS)':
                s.write('\t\tAetna Pricing: ' + item.get('Charge') + '\n')
            else:
                s.write('\t\tAetna Pricing: N/A \n')
                pass

            if item.get('Payer') == 'QUALCHOICE LIFE AND HEALTH INSURANCE COMPANY INC - (Indemnity)':
                s.write('\t\tQualChoice Pricing: ' + item.get('Charge') + '\n')
            else:        
                s.write('\t\tQualChoice Pricing: N/A \n\n')
                break

                
s.close()
```

```python
# COMBINATION (RUN THIS CODE!)
"""
This code makes it easy to compare the gross charges and Aetna/QualChoice pricing of common procedures between two hospitals: 
White River and Saline. Since this a comparison, I assume that one would only look at procedures that both hospitals have, meaning
if White River has a procedure that Saline does not, it will not be included in the file. This code produces a text file called 'stretch' that
shows for each procedure/code the gross charges, Aetna pricing, and QualChoice pricing of each hospital.
"""

import json
from lxml import objectify

with open('/data/whiteriver.json') as j: # open json file and close when done
    root = json.load(j)
with open('/data/saline.xml') as x: # open xml file and close when done
    sa = objectify.parse(x)

s = open('stretch2.txt', 'w') # open new text file called stretch and write into when called

main = sa.getroot() # get main root of saline file

for facility in main.findall('Facility'): # to be completely honest, I forgot why I have these two lines of code...
    name = facility.get('Name') # ...but if I try to remove them, my code doesn't work, so they get to stay


wr_charges = root['root']['StandardCharges'] # assign standard charges key to wr_charges

for wr_charge in wr_charges: # iterate through each charge found in whiteriver
    patients = facility.getchildren() # assign children of facility to patients

    for patient in patients: # iterate through each kind of patient type (inpatient, emergency, etc..)
        if patient.get('Type') == 'Outpatient': # if the type is Outpatient...
            sa_charges = patient.getchildren() # ...then assign its children to sa_charges

    for sa_charge in sa_charges: # iterate through each sa_charge
        codes = sa_charge.getchildren() # assign each child(ren) to codes

    for code in codes: # iterate through each code
        if code.get('Code') == wr_charge['ProcedureCode']: # if the current code is found in both whiteriver and saline files...
            s.write('Procedure Code: ' + wr_charge['ProcedureCode'] + '\n') # ...then write the procedure code into stretch.txt
            items = code.getchildren() # assign code children to items
            
            s.write('Description: ' + wr_charge['Description'] + '\n\n') # write description into stretch.txt
            
            s.write('\tWhite River Gross Charges: ' + wr_charge.get('OutpatientGrossCharge') + '\n') # write the gross charges for both white river and saline
            s.write('\tSaline Gross Charges: ' + str(items[2]) + '\n\n')
            
            for item in code.iter('Contract'): # for each item found within contract
                s.write('\tWhite River Aetna Pricing: ' + wr_charge.get('AETNA_Outpatient') + '\n') # write white river aetna pricing
                if item.get('Payer') == 'AETNA - (POS)': # determines if saline has aetna pricing
                    s.write('\tSaline Aetna Pricing: ' + item.get('Charge') + '\n\n') # write aetna pricing if it is found
                else:
                    s.write('\tSaline Aetna Pricing: N/A \n\n') # write n/a if it is not found
                    pass # skip to next payer

                s.write('\tWhite River QualChoice Pricing: ' + wr_charge.get('QUALCHOICE_Outpatient') + '\n') # write white river qualchoice pricing
                if item.get('Payer') == 'QUALCHOICE LIFE AND HEALTH INSURANCE COMPANY INC - (Indemnity)': # determines if saline has qualchoice pricing
                    s.write('\tSaline QualChoice Pricing: ' + item.get('Charge') + '\n\n\n') # write qualchoice pricing if it is found
                else:        
                    s.write('\tSaline QualChoice Pricing: N/A \n\n\n') # write n/a if it is not found
                    break # stop loop
    
s.close() # close stretch.txt

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.